<div class="alert alert-info">
    
➡️ Before you start, make sure that you are familiar with the **[study guide](https://liu-nlp.ai/text-mining/logistics/)**, in particular the rules around **cheating and plagiarism** (found in the course memo).

➡️ If you use code from external sources (e.g. StackOverflow, ChatGPT, ...) as part of your solutions, don't forget to add a reference to these source(s) (for example as a comment above your code).

➡️ Make sure you fill in all cells that say **`YOUR CODE HERE`** or **YOUR ANSWER HERE**.  You normally shouldn't need to modify any of the other cells.

</div>

# L1: Information Retrieval

In this lab you will apply basic techniques from information retrieval to implement the core of a minimalistic search engine. The data for this lab consists of a collection of app descriptions scraped from the [Google Play Store](https://play.google.com/store/apps?hl=en). From this collection, your search engine should retrieve those apps whose descriptions best match a given query under the vector space model.

```python
# Define some helper functions that are used in this notebook
from IPython.display import display, HTML

def success():
    display(HTML('<div class="alert alert-success"><strong>Checks have passed!</strong></div>'))
```

## Dataset

The app descriptions come in the form of a compressed [JSON](https://en.wikipedia.org/wiki/JSON) file. Start by loading this file into a [Pandas](https://pandas.pydata.org) [DataFrame](https://pandas.pydata.org/pandas-docs/stable/getting_started/dsintro.html#dataframe).

```python
import bz2
import numpy as np
import pandas as pd
pd.set_option('display.max_colwidth', 500)

with bz2.open('app-descriptions.json.bz2', mode='rt', encoding='utf-8') as source:
    df = pd.read_json(source, encoding='utf-8')
```

In Pandas, a DataFrame is a table with indexed rows and labelled columns of potentially different types. You can access data in a DataFrame in various ways, including by row and column. To give an example, the code in the next cell shows rows 200–204:

```python
df.loc[200:205]
```

Output:
```
name  \
200  Brick Breaker Star: Space King   
201      Brick Classic - Brick Game   
202     Bricks Breaker - Glow Balls   
203            Bricks Breaker Quest   
204             Brothers in Arms® 3   
205       Brown Dust - Tactical RPG   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                             description  
200  Introducing the best Brick Breaker game that everyone can enjoy.\nEnjoy various missions and addictively simple play control.\n\n[Features]\n- Hundreds of stages and various missions\n- No limit to play such as Heart, play as much as you can!\n-
```

As you can see, there are two labelled columns: `name` (the name of the app) and `description` (a textual description). The next cell shows how to access only the description field from row 200:

```python
df.loc[200, 'description']
```

Output:
```
'Introducing the best Brick Breaker game that everyone can enjoy.\nEnjoy various missions and addictively simple play control.\n\n[Features]\n- Hundreds of stages and various missions\n- No limit to play such as Heart, play as much as you can!\n- 5 kinds of various items and items reinforcement system\n- No network required\n- game file is as low as 20M, light-weight download!\n- supports tablet screen\n- supports Google Play Leaderboards, Achievement, Multiplay\n- supports 14 languages\n\nHomepage:\nhttps://play.google.com/store/apps/dev?id=4931745640662708567\n\nFacebook: \nhttps://www.facebook.com/spcomesgames/'
```

## Problem 1: What's in a vector?

We start by vectorising the data — more specifically, we map each app description to a tf–idf vector. This is very simple with a library like [scikit-learn](https://scikit-learn.org/stable/), which provides a [TfidfVectorizer](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html) class for exactly this purpose.  If we instantiate this class, and call `fit_transform()` on all of our app descriptions, scikit-learn will preprocess and tokenize each app description, compute tf–idf values for each of them, and return a vectorised representation:

```python
from sklearn.feature_extraction.text import TfidfVectorizer

vectorizer = TfidfVectorizer()
X = vectorizer.fit_transform(df['description'])
X
```

Output:
```
<Compressed Sparse Row sparse matrix of dtype 'float64'
	with 267110 stored elements and shape (1614, 27877)>
```

Let’s pick the app "Pancake Tower", which has a rather short description text, to see how it has been vectorised:

```python
# We can use 'toarray' to convert the sparse matrix object into a "normal" array
vec = X[1032].toarray()[0]

# The app description & its corresponding vector
df.loc[1032, 'description'], vec
```

Output:
```
("Let's see how many pancakes you can pile up!!",
 array([0., 0., 0., ..., 0., 0., 0.], shape=(27877,)))
```

That's not very informative yet.  We know that the vector contains tf–idf values, and that each dimension of the vector corresponds to a token in the vectorizer’s vocabulary; let's extract these for this specific example.

Your **first task** is to find out how to access the `vectorizer`’s vocabulary, for example by [checking the documentation of `TfidfVectorizer`](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html), and print all the tokens that are represented in the vector with a tf–idf value greater than zero (i.e., only the tokens that are actually part of this app’s description) _in descending order of the tf–idf values_.  In other words, the token with the highest tf–idf value should be at the top of your output, and the token with the lowest tf–idf value at the bottom.   Before you implement this, think about what you would expect the output look like, for example which words you would expect to have the highest/lowest tf–idf values in this example.

Your final output should look something like this:

```
<token 1>: <tf-idf value 1>
<token 2>: <tf-idf value 2>
...
```

<!-- grade_id: cell-95a3350b5c599c39, points: 0, grade: true, solution: true -->
```python
"""Print the tokens and their tf–idf values, in descending order."""

# YOUR CODE HERE
# encoded length is equal to the length of vocabulary
# but the value is encoded differently from one-hot, not only position info
# TF-IDF = TF × IDF
# TF: term frequency in present documentation, fre_in_doc(token, doc)/token_length(doc)
# IDF: term frequency in the documents, log((length(docs)+1) / length(sum(token in docs)) + 1)
# IDF is fixed but not IF

# sort vocab based on their id
# return a list consisting of tuples(token, token_id)
vocab = vectorizer.vocabulary_
sorted_vocab = sorted(vocab.items(), key=lambda x: x[1])
# get the non zero vector 
nonzero_vec_indices = np.where(vec > 0)[0]
nonzero_vec = vec[nonzero_vec_indices]
# get the corresponding tokens
nonzero_tokens = [sorted_vocab[index][0] for index in nonzero_vec_indices]
# create a vector to token dict and sort it
vec2token = {token: float(vector) for token, vector in zip(nonzero_tokens, nonzero_vec)}
sorted_items = sorted(vec2token.items(), key=lambda x: x[1], reverse=True)
sorted_items

```

Output:
```
[('pancakes', 0.6539332651185913),
 ('pile', 0.5304701435508047),
 ('let', 0.2615287714771797),
 ('see', 0.2557630827415271),
 ('many', 0.23491959669849022),
 ('how', 0.21153246225085887),
 ('up', 0.17216837691451817),
 ('can', 0.13047602895910532),
 ('you', 0.10276923239718011)]
```

## Problem 2: Finding the nearest vectors

To build a small search engine, we need to be able to turn _queries_ (for example the string "pile up pancakes") into _query vectors_, and then find out which of our app description vectors are closest to the query vector.

For the first part (turning queries into query vectors), we can simply re-use the `vectorizer` that we used for the app descriptions. For the second part, an easy way to find the closest vectors is to use scikit-learn’s [NearestNeighbors](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html) class. This class needs to be _fit_ on a set of vectors (the "training set"; in our case the app descriptions) and can then be used with any vector to find its _nearest neighbors_ in the vector space.

**First,** instantiate and fit a class that returns the _ten (10)_ nearest neighbors:

<!-- grade_id: cell-bcebb6d4a62e7b7f, points: 0, grade: false, solution: true -->
```python
"""Instantiate and fit a class that returns the 10 nearest neighboring vectors."""
# YOUR CODE HERE
from sklearn.neighbors import NearestNeighbors
n = 10

nn = NearestNeighbors(n_neighbors=n, metric='cosine')  
nn.fit(X)
```

Output:
```
NearestNeighbors(metric='cosine', n_neighbors=10)
```

**Second,** implement a function that uses the vectorizer and the fitted class to find the nearest neighbours for a given query string:

<!-- grade_id: cell-18bf238a099d709b, points: 0, grade: false, solution: true -->
```python
def search(query):
    """Find the nearest neighbors in `df` for a query string.

    Arguments:
      query (str): A query string.

    Returns:
      The 10 apps (with name and description) most similar (in terms of
      cosine similarity) to the given query as a Pandas DataFrame.
    """
    # YOUR CODE HERE
    # Iterable parameter required
    # why it would compress itself......
    encoded_query = vectorizer.transform([query]).toarray()
    # indices gets a 2d list, don't know why 
    _, indices = nn.kneighbors(encoded_query)
    indices = indices.ravel()
    result_pd = pd.DataFrame({"name": df['name'].loc[indices], "description": df['description'].loc[indices]})
    return result_pd.reset_index(drop=True)
```

### 🤞 Test your code

Test your implementation by running the following cell, which will sanity-check your return value and show the 10 best search results for the query _"pile up pancakes"_:

<!-- grade_id: cell-8506d7fe5961f87b, points: 1, grade: true, solution: false -->
```python
"""Check that searching for "pile up pancakes" returns a DataFrame with ten results,
   and that the top result is "Pancake Tower"."""

result = search('pile up pancakes')
display(result)
assert isinstance(result, pd.DataFrame), "search() function should return a Pandas DataFrame"
assert len(result) == 10, "search() function should return 10 search results"
assert result.iloc[0]["name"] == "Pancake Tower", "Top search result should be 'Pancake Tower'"
success()
```

Output:
```
name  \
0                                      Pancake Tower   
1                    Cooking School: Games for Girls   
2  Hell’s Cooking — crazy chef burger, kitchen fever   
3                                          Solitaire   
4                                       Rummy - Free   
5                       Sago Mini Trucks and Diggers   
6                        Dr. Panda's Ice Cream Truck   
7                                    Turbo Dismount™   
8                                              UNO!™   
9                                     TO-FU Oh!SUSHI   

                                                                                                                                                                                                                                                                                                                                                                                                                                               
```

Before continuing with the next problem, play around a bit with this simple search functionality by trying out different search queries, and see if the results look like what you would expect:

```python
# Example — try out your own queries!
search("dodge trains")
```

Output:
```
name  \
0                               Train Conductor World   
1                                      Subway Surfers   
2                              Subway Princess Runner   
3                      No Humanity - The Hardest Game   
4                                          Bus Rush 2   
5                     Virus War - Space Shooting Game   
6                       Dancing Road: Color Ball Run!   
7  Bob - jigsaw puzzles free games for kids & parents   
8                      Blocky Highway: Traffic Racing   
9                           Cat Runner: Decorate Home   

                                                                                                                                                                                                                                                                                                                                                                                                                                     
```

## Problem 3: Custom preprocessing & tokenization

In Problem 1, you should have seen that `TfidfVectorizer` already performs some preprocessing by default and also does its own tokenization of the input data. This is great for getting started, but often we want to have more control over these steps. We can customize some aspects of the preprocessing through arguments when instantiating `TfidfVectorizer`, but for this exercise, we want to do _all_ of our preprocessing & tokenizing outside of scikit-learn.

Concretely, we want to use [spaCy](https://spacy.io), a library that we will make use of in later labs as well.  Here is a brief example of how to load and use a spaCy model:

```python
import spacy
# Load the small English model, disabling some components that we don't need right now
nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])

# Take an example sentence and print every token from it separately
doc = nlp("Apple is looking at buying U.K. startup for $1 billion")
for token in doc:
    print(token.text)
```

Output:
```
Apple
is
looking
at
buying
U.K.
startup
for
$
1
billion
```

**Your task** is to write a preprocessing function that uses spaCy to perform the following steps:
- tokenization
- lemmatization
- stop word removal
- removing tokens containing non-alphabetical characters

We recommend that you go through the [Linguistic annotations](https://spacy.io/usage/spacy-101#annotations) section of the spaCy&nbsp;101, which demonstrates how you can get the relevant kind of information via the spaCy library.

Implement your preprocessor by completing the following function:

<!-- grade_id: cell-2df4eac94cdf6be3, points: 0, grade: false, solution: true -->
```python
def preprocess(text):
    """Preprocess the given text by tokenising it, removing any stop words, 
    replacing each remaining token with its lemma (base form), and discarding 
    all lemmas that contain non-alphabetical characters.

    Arguments:
      text (str): The text to preprocess.

    Returns:
      The list of remaining lemmas after preprocessing (represented as strings).
    """
    # YOUR CODE HERE
    # text to tokens
    tokens = nlp(text)
    # the functions below cannot apply on str
    # print(type(tokens[0]))
    # lemma_ -> lemmatization
    # is_stop -> stop word removal
    # is_alpha -> remove non-letter words
    processed_tokens = [token.lemma_ for token in tokens if token.is_alpha and not token.is_stop]
    return processed_tokens
  
```

### 🤞 Test your code

Test your implementation by running the following cell:

<!-- grade_id: cell-642185b139f2cef6, points: 1, grade: true, solution: false -->
```python
"""Check that the preprocessing returns the correct output for a number of test cases."""

assert (
    preprocess('Apple is looking at buying U.K. startup for $1 billion') ==
    ['Apple', 'look', 'buy', 'startup', 'billion']
)
assert (
    preprocess('"Love Story" is a country pop song written and sung by Taylor Swift.') ==
    ['Love', 'Story', 'country', 'pop', 'song', 'write', 'sing', 'Taylor', 'Swift']
)
success()
```

Output:
```
<IPython.core.display.HTML object>
```

## Problem 4: The effect of preprocessing

To make use of the new `preprocess` function from Problem 3, we need to make sure that we incorporate it into `TfidfVectorizer` and disable all preprocessing & tokenization that `TfidfVectorizer` performs by default. Afterwards, we also need to re-fit the vectorizer and the nearest-neighbors class. To make this a bit easier to handle, let’s take everything we have done so far and put it in a single class `AppSearcher`.

### Task 4.1

**Your first task** is to complete the stub of the `AppSearcher` class given below. Keep in mind:
- The `fit()` function should fit both the vectorizer (from Problem 1) and the nearest-neighbors class (from Problem 2).  Make sure to modify the call to `TfidfVectorizer` to _disable all preprocessing & tokenization_ that it would do by default, and replace it with a call to the `preprocess()` function _defined in `AppSearcher`_.
- For the `preprocess()` function, you can start by copying your solution from Problem 3.
- For the `search()` function, you can copy your solution from Problem 2.
- Make sure to adapt your code to store the everything (data, vectorizer, nearest-neighbors class) within the `AppSearcher` class, so that your solution is independent of the code you wrote above!

<!-- grade_id: cell-cb53b61e9efe98af, points: 0, grade: true, solution: true -->
```python
class AppSearcher:
    def fit(self, df, reults_num=10):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        # fit vectorizer
        vectorizer = TfidfVectorizer()
        feature_matrix = vectorizer.fit_transform(df['description'])
        self.vectorizer = vectorizer
        # fit nearest neighbor
        nn = NearestNeighbors(n_neighbors=reults_num, metric='cosine')  
        nn.fit(feature_matrix)
        self.nn = nn

    def preprocess(self, text, lemma=True, lowercase=True, stopwords_removal=True, nonletter_words=True):
        """
        Preprocess the given text (cf. Problem 3).
        Tasks include:
            -  
            - lowercasing all characters
            - removing stop words
            - removing tokens containing non-alphabetical characters
        """
        # YOUR CODE HERE
        # load the model
        nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        doc = nlp(text)
        # to prevent the type(outcome) from last round becomes str
        # cope each token with full process
        tokens = []
        for token in doc:
            if nonletter_words and not token.is_alpha:
                continue
            if stopwords_removal and token.is_stop:
                continue

            tok = token.lemma_ if lemma else token.text
            if lowercase:
                tok = tok.lower()

            tokens.append(tok)

        result = " ".join(tokens)
        return result

    def search(self, query, lemma=True, lowercase=True, stopwords_removal=True, nonletter_words=True):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        tokens = self.preprocess(query, lemma, lowercase, stopwords_removal, nonletter_words)
        encoded_query = self.vectorizer.transform([tokens]).toarray()
        _, indices = self.nn.kneighbors(encoded_query)
        indices = indices.ravel()
        result_pd = pd.DataFrame({"name": df['name'].loc[indices], "description": df['description'].loc[indices]})
        return result_pd.reset_index(drop=True)

```

#### 🤞 Test your code

The following cell demonstrates how your class should be used. Note that it can take a bit longer to train it on the data as before, since we’re now calling spaCy for the preprocessing.

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  \
0                   Pancake Tower   
1                       Solitaire   
2                    Rummy - Free   
3    Sago Mini Trucks and Diggers   
4                 Turbo Dismount™   
5     Dr. Panda's Ice Cream Truck   
6                  Solitaire Free   
7  Dr. Panda Ice Cream Truck Free   
8                  TO-FU Oh!SUSHI   
9                           UNO!™   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           description  
0                                                                                                                   
```

### Task 4.2

**Your second task** is to experiment with the effect of using (or not using) different preprocessing steps.  We always need to _tokenize_ the text, but other preprocessing steps are optional and require a conscious decision whether to use them or not, such as:
- lemmatization
- lowercasing all characters
- removing stop words
- removing tokens containing non-alphabetical characters

**Modify the definition of the `preprocess()` function** of `AppSearcher` to include/exclude individual preprocessing steps, run some searches, and observe if and how the results change.  Which search queries you try out is up to you — you could compare searching for "pile up pancakes" with "pancake piling", for example; or you could try entirely different search queries aimed at different kinds of apps.  (You can modify the class directly by changing the cell above under Task 4.1, or copy the definitions to the cells below, whichever you prefer; there is no separate code to show for this task, but you will use your observations here for the individual reflection.)

```python
# if I use the previous query, whose reult is existed in the list
# probably hardy detect difference 
query = "Track the Best AI-based Budget planner for Students v2.0!"
apps.fit(df, reults_num=5)
apps.search(query, lemma=True, lowercase=True, stopwords_removal=True, nonletter_words=True)

```

Output:
```
name  \
0  Lifesum - Diet Plan, Macro Calculator & Food Diary   
1                            Instant Match For Tinder   
2                                           2D Strike   
3                      Complete Anatomy Platform 2020   
4                                       Durak | Дурак   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           description  
0  Diet plan, food diary, macro calculator, calorie counter & healthy recipes, all in one convenient place. Reach your goals with food tracker Lifesum!\n\nKeto diet, high protein or maybe vegan? We’ll 
```

```python
apps.search(query, lemma=False, lowercase=True, stopwords_removal=True, nonletter_words=True)

```

Output:
```
name  \
0                                       Durak | Дурак   
1  Lifesum - Diet Plan, Macro Calculator & Food Diary   
2                            Instant Match For Tinder   
3              Todoist: To-Do List, Tasks & Reminders   
4                                    Tic Tac Toe Glow   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           description  
0  "Durak" or "Дурак" is the most popular Russian card game. This game allows to play with 1 to 3 AI players. It's also been know as "v duraka" or " В дурака". \n\nFeatures:\n- sounds\n- vibration (e.g
```

```python
apps.search(query, lemma=True, lowercase=False, stopwords_removal=True, nonletter_words=True)

```

Output:
```
name  \
0  Lifesum - Diet Plan, Macro Calculator & Food Diary   
1                            Instant Match For Tinder   
2                                           2D Strike   
3                      Complete Anatomy Platform 2020   
4                                       Durak | Дурак   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           description  
0  Diet plan, food diary, macro calculator, calorie counter & healthy recipes, all in one convenient place. Reach your goals with food tracker Lifesum!\n\nKeto diet, high protein or maybe vegan? We’ll 
```

```python
apps.search(query, lemma=True, lowercase=True, stopwords_removal=False, nonletter_words=True)

```

Output:
```
name  \
0                              Last Shelter: Survival   
1                                       World on Fire   
2                      Complete Anatomy Platform 2020   
3  Lifesum - Diet Plan, Macro Calculator & Food Diary   
4                        Rocket War: Clash in the Fog   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           description  
0  Tons of Rewards, Limited Packages, Don't wait, ACT!\n\nThey’re Here…..\n\nYour Mission, Survive\n\nThe war ravages\nThe virus has spread on a global scale, and most of the humans are now zombies.In 
```

```python
apps.search(query, lemma=True, lowercase=True, stopwords_removal=True, nonletter_words=False)
```

Output:
```
name  \
0  Lifesum - Diet Plan, Macro Calculator & Food Diary   
1                            Instant Match For Tinder   
2                                           2D Strike   
3                      Complete Anatomy Platform 2020   
4                                       Durak | Дурак   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                           description  
0  Diet plan, food diary, macro calculator, calorie counter & healthy recipes, all in one convenient place. Reach your goals with food tracker Lifesum!\n\nKeto diet, high protein or maybe vegan? We’ll 
```

```python
# code here for supporting reflection
# question 1
# for description: "Let's see how many pancakes you can pile up!!"
max_ = np.argmax(vec)
min_ = np.argmin(vec)

max_token = sorted_vocab[max_][0]
min_token = sorted_vocab[min_][0]

print(max_token)
print(min_token)


```

Output:
```
pancakes
00
tsuki
applikation
```

## Individual reflection

<div class="alert alert-info">
    <strong>After you have solved the lab,</strong> write a <em>brief</em> reflection (max. one A4 page) on the question(s) below.  Remember:
    <ul>
        <li>You are encouraged to discuss this part with your lab partner, but you should each write up your reflection <strong>individually</strong>.</li>
        <li><strong>Do not put your answers in the notebook</strong>; upload them in the separate submission opportunity for the reflections on Lisam.</li>
    </ul>
</div>

1. In Problem 1, which token had the highest tf–idf score, which the lowest?  Based on your knowledge of how tf–idf works, how would you explain this result?
2. Based on your observations in Problem 4, which preprocessing steps do you think are the most appropriate for this "search engine" example?  Why?

**Congratulations on finishing this lab! 👍**

<div class="alert alert-info">
    
➡️ Before you submit, **make sure the notebook can be run from start to finish** without errors.  For this, _restart the kernel_ and _run all cells_ from top to bottom. In Jupyter Notebook version 7 or higher, you can do this via "Run$\rightarrow$Restart Kernel and Run All Cells..." in the menu (or the "⏩" button in the toolbar).

</div>