<div class="alert alert-info">
    
➡️ Before you start, make sure that you are familiar with the **[study guide](https://liu-nlp.ai/text-mining/logistics/)**, in particular the rules around **cheating and plagiarism** (found in the course memo).

➡️ If you use code from external sources (e.g. StackOverflow, ChatGPT, ...) as part of your solutions, don't forget to add a reference to these source(s) (for example as a comment above your code).

➡️ Make sure you fill in all cells that say **`YOUR CODE HERE`** or **YOUR ANSWER HERE**.  You normally shouldn't need to modify any of the other cells.

</div>

# L1: Information Retrieval

In this lab you will apply basic techniques from information retrieval to implement the core of a minimalistic search engine. The data for this lab consists of a collection of app descriptions scraped from the [Google Play Store](https://play.google.com/store/apps?hl=en). From this collection, your search engine should retrieve those apps whose descriptions best match a given query under the vector space model.

```python
# Define some helper functions that are used in this notebook
from IPython.display import display, HTML

def success():
    display(HTML('<div class="alert alert-success"><strong>Checks have passed!</strong></div>'))
```

## Dataset

The app descriptions come in the form of a compressed [JSON](https://en.wikipedia.org/wiki/JSON) file. Start by loading this file into a [Pandas](https://pandas.pydata.org) [DataFrame](https://pandas.pydata.org/pandas-docs/stable/getting_started/dsintro.html#dataframe).

```python
import bz2
import numpy as np
import pandas as pd
pd.set_option('display.max_colwidth', 500)

with bz2.open('app-descriptions.json.bz2', mode='rt', encoding='utf-8') as source:
    df = pd.read_json(source, encoding='utf-8')
```

In Pandas, a DataFrame is a table with indexed rows and labelled columns of potentially different types. You can access data in a DataFrame in various ways, including by row and column. To give an example, the code in the next cell shows rows 200–204:

```python
df.loc[200:205]
```

Output:
```
name  \
200  Brick Breaker Star: Space King   
201      Brick Classic - Brick Game   
202     Bricks Breaker - Glow Balls   
203            Bricks Breaker Quest   
204             Brothers in Arms® 3   
205       Brown Dust - Tactical RPG   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                             description  
200  Introducing the best Brick Breaker game that everyone can enjoy.\nEnjoy various missions and addictively simple play control.\n\n[Features]\n- Hundreds of stages and various missions\n- No limit to play such as Heart, play as much as you can!\n-
```

As you can see, there are two labelled columns: `name` (the name of the app) and `description` (a textual description). The next cell shows how to access only the description field from row 200:

```python
df.loc[200, 'description']
```

Output:
```
'Introducing the best Brick Breaker game that everyone can enjoy.\nEnjoy various missions and addictively simple play control.\n\n[Features]\n- Hundreds of stages and various missions\n- No limit to play such as Heart, play as much as you can!\n- 5 kinds of various items and items reinforcement system\n- No network required\n- game file is as low as 20M, light-weight download!\n- supports tablet screen\n- supports Google Play Leaderboards, Achievement, Multiplay\n- supports 14 languages\n\nHomepage:\nhttps://play.google.com/store/apps/dev?id=4931745640662708567\n\nFacebook: \nhttps://www.facebook.com/spcomesgames/'
```

## Problem 1: What's in a vector?

We start by vectorising the data — more specifically, we map each app description to a tf–idf vector. This is very simple with a library like [scikit-learn](https://scikit-learn.org/stable/), which provides a [TfidfVectorizer](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html) class for exactly this purpose.  If we instantiate this class, and call `fit_transform()` on all of our app descriptions, scikit-learn will preprocess and tokenize each app description, compute tf–idf values for each of them, and return a vectorised representation:

```python
from sklearn.feature_extraction.text import TfidfVectorizer

vectorizer = TfidfVectorizer()
X = vectorizer.fit_transform(df['description'])
X
```

Output:
```
<Compressed Sparse Row sparse matrix of dtype 'float64'
	with 267110 stored elements and shape (1614, 27877)>
```

Let’s pick the app "Pancake Tower", which has a rather short description text, to see how it has been vectorised:

```python
# We can use 'toarray' to convert the sparse matrix object into a "normal" array
vec = X[1032].toarray()[0]

# The app description & its corresponding vector
df.loc[1032, 'description'], vec
```

Output:
```
("Let's see how many pancakes you can pile up!!",
 array([0., 0., 0., ..., 0., 0., 0.], shape=(27877,)))
```

That's not very informative yet.  We know that the vector contains tf–idf values, and that each dimension of the vector corresponds to a token in the vectorizer’s vocabulary; let's extract these for this specific example.

Your **first task** is to find out how to access the `vectorizer`’s vocabulary, for example by [checking the documentation of `TfidfVectorizer`](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html), and print all the tokens that are represented in the vector with a tf–idf value greater than zero (i.e., only the tokens that are actually part of this app’s description) _in descending order of the tf–idf values_.  In other words, the token with the highest tf–idf value should be at the top of your output, and the token with the lowest tf–idf value at the bottom.   Before you implement this, think about what you would expect the output look like, for example which words you would expect to have the highest/lowest tf–idf values in this example.

Your final output should look something like this:

```
<token 1>: <tf-idf value 1>
<token 2>: <tf-idf value 2>
...
```

<!-- grade_id: cell-95a3350b5c599c39, points: 0, grade: true, solution: true -->
```python
"""Print the tokens and their tf–idf values, in descending order."""

# Get the tf-idf vector for document with index 1032
vec = X[1032].toarray()[0]
def sorted_tokens(vectorizer, vec):
    """Print the tokens and their tf–idf values, in descending order."""
    idf = vectorizer.idf_
    voca = vectorizer.vocabulary_
    # Get feature names and values for non-zero entries
    feature_names = vectorizer.get_feature_names_out()
    nonzero_idx = vec.nonzero()[0]
    tokens = [(feature_names[i], vec[i]) for i in nonzero_idx]
    # Sort by tf-idf value in descending order and print
    tokens_sorted = sorted(tokens, key=lambda x: x[1], reverse=True)
    for tok, val in tokens_sorted:
        print(f"{tok}: {val}")
sorted_tokens(vectorizer, vec)
```

Output:
```
pancakes: 0.6539332651185913
pile: 0.5304701435508047
let: 0.2615287714771797
see: 0.2557630827415271
many: 0.23491959669849022
how: 0.21153246225085887
up: 0.17216837691451817
can: 0.13047602895910532
you: 0.10276923239718011
```

## Problem 2: Finding the nearest vectors

To build a small search engine, we need to be able to turn _queries_ (for example the string "pile up pancakes") into _query vectors_, and then find out which of our app description vectors are closest to the query vector.

For the first part (turning queries into query vectors), we can simply re-use the `vectorizer` that we used for the app descriptions. For the second part, an easy way to find the closest vectors is to use scikit-learn’s [NearestNeighbors](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html) class. This class needs to be _fit_ on a set of vectors (the "training set"; in our case the app descriptions) and can then be used with any vector to find its _nearest neighbors_ in the vector space.

**First,** instantiate and fit a class that returns the _ten (10)_ nearest neighbors:

<!-- grade_id: cell-bcebb6d4a62e7b7f, points: 0, grade: false, solution: true -->
```python
"""Instantiate and fit a class that returns the 10 nearest neighboring vectors."""

from sklearn.neighbors import NearestNeighbors
# Fit a NearestNeighbors model on the tf-idf matrix X
nn = NearestNeighbors(n_neighbors=10)
nn.fit(X)
```

Output:
```
NearestNeighbors(n_neighbors=10)
```

**Second,** implement a function that uses the vectorizer and the fitted class to find the nearest neighbours for a given query string:

<!-- grade_id: cell-18bf238a099d709b, points: 0, grade: false, solution: true -->
```python
def search(query):
    """Find the nearest neighbors in `df` for a query string.

    Arguments:
      query (str): A query string.

    Returns:
      The 10 apps (with name and description) most similar (in terms of
      cosine similarity) to the given query as a Pandas DataFrame.
    """
    # Transform the query to a tf-idf vector and find nearest neighbors
    qv = vectorizer.transform([query])
    distances, indices = nn.kneighbors(qv)
    res = df.iloc[indices[0]][['name', 'description']].copy()
    res['score'] = distances[0]
    return res.reset_index(drop=True)
```

### 🤞 Test your code

Test your implementation by running the following cell, which will sanity-check your return value and show the 10 best search results for the query _"pile up pancakes"_:

<!-- grade_id: cell-8506d7fe5961f87b, points: 1, grade: true, solution: false -->
```python
"""Check that searching for "pile up pancakes" returns a DataFrame with ten results,
   and that the top result is "Pancake Tower"."""

result = search('pile up pancakes')
display(result)
assert isinstance(result, pd.DataFrame), "search() function should return a Pandas DataFrame"
assert len(result) == 10, "search() function should return 10 search results"
assert result.iloc[0]["name"] == "Pancake Tower", "Top search result should be 'Pancake Tower'"
success()
```

Output:
```
name  \
0                                      Pancake Tower   
1                    Cooking School: Games for Girls   
2  Hell’s Cooking — crazy chef burger, kitchen fever   
3                                          Solitaire   
4                                       Rummy - Free   
5                       Sago Mini Trucks and Diggers   
6                        Dr. Panda's Ice Cream Truck   
7                                    Turbo Dismount™   
8                                              UNO!™   
9                                     TO-FU Oh!SUSHI   

                                                                                                                                                                                                                                                                                                                                                                                                                                               
```

Before continuing with the next problem, play around a bit with this simple search functionality by trying out different search queries, and see if the results look like what you would expect:

```python
# Example — try out your own queries!
search("dodge trains")
```

Output:
```
name  \
0                               Train Conductor World   
1                                      Subway Surfers   
2                              Subway Princess Runner   
3                      No Humanity - The Hardest Game   
4                                          Bus Rush 2   
5                     Virus War - Space Shooting Game   
6                       Dancing Road: Color Ball Run!   
7  Bob - jigsaw puzzles free games for kids & parents   
8                      Blocky Highway: Traffic Racing   
9                           Cat Runner: Decorate Home   

                                                                                                                                                                                                                                                                                                                                                                                                                                     
```

## Problem 3: Custom preprocessing & tokenization

In Problem 1, you should have seen that `TfidfVectorizer` already performs some preprocessing by default and also does its own tokenization of the input data. This is great for getting started, but often we want to have more control over these steps. We can customize some aspects of the preprocessing through arguments when instantiating `TfidfVectorizer`, but for this exercise, we want to do _all_ of our preprocessing & tokenizing outside of scikit-learn.

Concretely, we want to use [spaCy](https://spacy.io), a library that we will make use of in later labs as well.  Here is a brief example of how to load and use a spaCy model:

```python
import spacy
# Load the small English model, disabling some components that we don't need right now
nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])

# Take an example sentence and print every token from it separately
doc = nlp("Apple is looking at buying U.K. startup for $1 billion")
for token in doc:
    print(token.text)
```

Output:
```
Apple
is
looking
at
buying
U.K.
startup
for
$
1
billion
```

**Your task** is to write a preprocessing function that uses spaCy to perform the following steps:
- tokenization
- lemmatization
- stop word removal
- removing tokens containing non-alphabetical characters

We recommend that you go through the [Linguistic annotations](https://spacy.io/usage/spacy-101#annotations) section of the spaCy&nbsp;101, which demonstrates how you can get the relevant kind of information via the spaCy library.

Implement your preprocessor by completing the following function:

<!-- grade_id: cell-2df4eac94cdf6be3, points: 0, grade: false, solution: true -->
```python
def preprocess(text):
    """Preprocess the given text by tokenising it, removing any stop words, 
    replacing each remaining token with its lemma (base form), and discarding 
    all lemmas that contain non-alphabetical characters. Proper nouns (PROPN)
    are kept in their original surface form and capitalization to preserve names."""
    doc = nlp(text)
    lemmas = []
    for token in doc:
        # skip stop words and non-alphabetic tokens
        if token.is_stop:
            continue
        if not token.is_alpha:
            continue
        if token.pos_ == 'PROPN': # if it is a proper noun, such as a name, keep original form  
            lemmas.append(token.text)
        else:
            lemmas.append(token.lemma_.lower())
    return lemmas
```

### 🤞 Test your code

Test your implementation by running the following cell:

<!-- grade_id: cell-642185b139f2cef6, points: 1, grade: true, solution: false -->
```python
"""Check that the preprocessing returns the correct output for a number of test cases."""

assert (
    preprocess('Apple is looking at buying U.K. startup for $1 billion') ==
    ['Apple', 'look', 'buy', 'startup', 'billion']
)
assert (
    preprocess('"Love Story" is a country pop song written and sung by Taylor Swift.') ==
    ['Love', 'Story', 'country', 'pop', 'song', 'write', 'sing', 'Taylor', 'Swift']
)
success()
```

Output:
```
<IPython.core.display.HTML object>
```

## Problem 4: The effect of preprocessing

To make use of the new `preprocess` function from Problem 3, we need to make sure that we incorporate it into `TfidfVectorizer` and disable all preprocessing & tokenization that `TfidfVectorizer` performs by default. Afterwards, we also need to re-fit the vectorizer and the nearest-neighbors class. To make this a bit easier to handle, let’s take everything we have done so far and put it in a single class `AppSearcher`.

### Task 4.1

**Your first task** is to complete the stub of the `AppSearcher` class given below. Keep in mind:
- The `fit()` function should fit both the vectorizer (from Problem 1) and the nearest-neighbors class (from Problem 2).  Make sure to modify the call to `TfidfVectorizer` to _disable all preprocessing & tokenization_ that it would do by default, and replace it with a call to the `preprocess()` function _defined in `AppSearcher`_.
- For the `preprocess()` function, you can start by copying your solution from Problem 3.
- For the `search()` function, you can copy your solution from Problem 2.
- Make sure to adapt your code to store the everything (data, vectorizer, nearest-neighbors class) within the `AppSearcher` class, so that your solution is independent of the code you wrote above!

<!-- grade_id: cell-cb53b61e9efe98af, points: 0, grade: true, solution: true -->
```python
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        import spacy
        from sklearn.feature_extraction.text import TfidfVectorizer
        from sklearn.neighbors import NearestNeighbors
        self.df = df
        # load spacy model once and store it
        self.nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])
        # Use custom tokenizer (self.preprocess) and disable TfidfVectorizer's own preprocessing
        self.vectorizer = TfidfVectorizer(preprocessor=lambda x: x, tokenizer=self.preprocess, lowercase=False)
        self.X = self.vectorizer.fit_transform(self.df['description'])
        self.nn = NearestNeighbors(n_neighbors=10, metric='cosine')
        self.nn.fit(self.X)

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # Use the stored spacy pipeline
        doc = self.nlp(text)
        tokens = []
        for token in doc:
            if token.is_stop:
                continue
            if not token.is_alpha:
                continue
            if token.pos_ == 'PROPN':
                tokens.append(token.text)
            else:
                tokens.append(token.lemma_.lower())
        return tokens

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        qv = self.vectorizer.transform([query])
        distances, indices = self.nn.kneighbors(qv)
        scores = 1 - distances[0]
        res = self.df.iloc[indices[0]][['name', 'description']].copy()
        res['score'] = scores
        return res.reset_index(drop=True)
```

#### 🤞 Test your code

The following cell demonstrates how your class should be used. Note that it can take a bit longer to train it on the data as before, since we’re now calling spaCy for the preprocessing.

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
/opt/anaconda3/envs/textmining/lib/python3.10/site-packages/sklearn/feature_extraction/text.py:517: UserWarning: The parameter 'token_pattern' will not be used since 'tokenizer' is not None'
  warnings.warn(
                                                name  \
0                                      Pancake Tower   
1                    Cooking School: Games for Girls   
2                                          Solitaire   
3                       Sago Mini Trucks and Diggers   
4                                   Spider Solitaire   
5  Hell’s Cooking — crazy chef burger, kitchen fever   
6                                       Rummy - Free   
7                        Dr. Panda's Ice Cream Truck   
8                                     Solitaire Free   
9                                    Turbo Dismount™   

                                                                                                                                                                               
```

### Task 4.2

**Your second task** is to experiment with the effect of using (or not using) different preprocessing steps.  We always need to _tokenize_ the text, but other preprocessing steps are optional and require a conscious decision whether to use them or not, such as:
- lemmatization
- lowercasing all characters
- removing stop words
- removing tokens containing non-alphabetical characters

**Modify the definition of the `preprocess()` function** of `AppSearcher` to include/exclude individual preprocessing steps, run some searches, and observe if and how the results change.  Which search queries you try out is up to you — you could compare searching for "pile up pancakes" with "pancake piling", for example; or you could try entirely different search queries aimed at different kinds of apps.  (You can modify the class directly by changing the cell above under Task 4.1, or copy the definitions to the cells below, whichever you prefer; there is no separate code to show for this task, but you will use your observations here for the individual reflection.)

```python
import spacy
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.neighbors import NearestNeighbors
# load a local spacy pipeline for these experiments (cached if already loaded)
nlp_local = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])

def make_tokenizer(lemmatize=True, lowercase=True, remove_stop=True, remove_nonalpha=True):
    def tokenizer(text):
        doc = nlp_local(text)
        toks = []
        for t in doc:
            if remove_stop and t.is_stop:
                continue
            if remove_nonalpha and not t.is_alpha:
                continue
            if t.pos_ == 'PROPN':
                tok = t.text
            else:
                tok = t.lemma_ if lemmatize else t.text
                if lowercase:
                    tok = tok.lower()
            toks.append(tok)
        return toks
    return tokenizer

def run_experiment(query='pile up pancakes', **flags):
    tok = make_tokenizer(**flags)
    vect = TfidfVectorizer(preprocessor=lambda x: x, tokenizer=tok, lowercase=False)
    Xexp = vect.fit_transform(df['description'])
    nnexp = NearestNeighbors(n_neighbors=20, metric='cosine').fit(Xexp)
    qv = vect.transform([query])

    # Show a small vocab sample
    print('vocab sample:', list(vect.get_feature_names_out())[:20])

    # Print query token scores (non-zero tf-idf entries) sorted by score
    # feature_names = vect.get_feature_names_out()
    # q_tokens = [(feature_names[i], qvec[i]) for i in qvec.nonzero()[0]]
    # q_tokens_sorted = sorted(q_tokens, key=lambda x: x[1], reverse=True)
    # print('query token scores:')
    # for tname, val in q_tokens_sorted:
    #     print(f'{tname}: {val:.6f}')

    # Also print token scores using helper if available (for compatibility)


    distances, indices = nnexp.kneighbors(qv)
    scores = 1 - distances[0]
    res = df.iloc[indices[0]][['name', 'description']].copy()
    res['score'] = scores
    display(res)
    return Xexp,vect


# Run a few variations (these re-fit the vectorizer each time)
Xexp1,vect1 = run_experiment('pile up pancakes', lemmatize=True, lowercase=True, remove_stop=True, remove_nonalpha=True)
Xexp2,vect2 =run_experiment('pile up pancakes', lemmatize=True, lowercase=True, remove_stop=True, remove_nonalpha=False)
Xexp3,vect3=run_experiment('pile up pancakes', lemmatize=True, lowercase=True, remove_stop=False, remove_nonalpha=True)
Xexp4,vect4 =run_experiment('pile up pancakes', lemmatize=True, lowercase=False, remove_stop=True, remove_nonalpha=True)
Xexp5,vect5 =run_experiment('pile up pancakes', lemmatize=False, lowercase=True, remove_stop=True, remove_nonalpha=True)
```

Output:
```
/opt/anaconda3/envs/textmining/lib/python3.10/site-packages/sklearn/feature_extraction/text.py:517: UserWarning: The parameter 'token_pattern' will not be used since 'tokenizer' is not None'
  warnings.warn(
vocab sample: ['AA', 'AAA', 'AAC', 'AAGBI', 'AAT', 'AB', 'ABC', 'ABCD', 'ABCmouse', 'ABCs', 'ABEL', 'ABG', 'ABI', 'ABS', 'ABSOLUTELY', 'AC', 'ACA', 'ACADEMY', 'ACHIEVEMENT', 'ACLS']
                                                   name  \
1032                                      Pancake Tower   
326                     Cooking School: Games for Girls   
1235                                          Solitaire   
1181                       Sago Mini Trucks and Diggers   
1263                                   Spider Solitaire   
656   Hell’s Cooking — crazy chef burger, kitchen fever   
1164                                       Rummy - Free   
436                         Dr. Panda's Ice Cream Truck   
1245                                     Solitaire Free   
1442                
```

```python

```

## Individual reflection

<div class="alert alert-info">
    <strong>After you have solved the lab,</strong> write a <em>brief</em> reflection (max. one A4 page) on the question(s) below.  Remember:
    <ul>
        <li>You are encouraged to discuss this part with your lab partner, but you should each write up your reflection <strong>individually</strong>.</li>
        <li><strong>Do not put your answers in the notebook</strong>; upload them in the separate submission opportunity for the reflections on Lisam.</li>
    </ul>
</div>

1. In Problem 1, which token had the highest tf–idf score, which the lowest?  Based on your knowledge of how tf–idf works, how would you explain this result?
2. Based on your observations in Problem 4, which preprocessing steps do you think are the most appropriate for this "search engine" example?  Why?

**Congratulations on finishing this lab! 👍**

<div class="alert alert-info">
    
➡️ Before you submit, **make sure the notebook can be run from start to finish** without errors.  For this, _restart the kernel_ and _run all cells_ from top to bottom. In Jupyter Notebook version 7 or higher, you can do this via "Run$\rightarrow$Restart Kernel and Run All Cells..." in the menu (or the "⏩" button in the toolbar).

</div>

```python
Y = Xexp1[656]
Y_ = Y.toarray()[0]
sorted_tokens(vect1,Y_)
```

Output:
```
burger: 0.38946512226495467
chef: 0.32588773772090207
kitchen: 0.2405834747191089
fever: 0.22667690449790226
hamburger: 0.2240581409086694
restaurant: 0.2033978950161184
cafe: 0.19473256113247733
Hell: 0.17924651272693554
diner: 0.17353320467919559
Cooking: 0.1471711033876947
girl: 0.14411168675965014
serve: 0.13959598980158408
crazy: 0.13338209120755382
cooking: 0.13337078055872112
visitor: 0.12112451771529135
hell: 0.1150093143379695
food: 0.09924597905367138
kitchenette: 0.09849683425007463
dash: 0.09849104129468418
tasty: 0.09639410910714163
ingredient: 0.09450378405547079
fast: 0.09236009306785811
cook: 0.09070497582834393
pancake: 0.08962325636346777
frenzy: 0.08676660233959779
game: 0.08436857869537219
coffee: 0.07667287622531299
craze: 0.07667287622531299
queen: 0.07453427208306297
prepare: 0.07422461553225629
guest: 0.07270231414209745
prepared: 0.07109999225623523
bake: 0.07036825785507392
meal: 0.06967609648947963
order: 0.06588986130332049
rush: 0.06382873625549057
real: 0.
```

```python
Y = Xexp5[656]
Y_ = Y.toarray()[0]
sorted_tokens(vect5,Y_)
```

Output:
```
burger: 0.339918401332665
chef: 0.3279992254143046
kitchen: 0.24478587803829457
diner: 0.23477667948243566
fever: 0.2281455866200747
cafe: 0.199388146094794
hamburger: 0.18782134358594854
restaurant: 0.1808544982719971
Hell: 0.1804078844568298
cooking: 0.15753560199151306
Cooking: 0.14812465253256532
crazy: 0.13569971963156507
girls: 0.1270583917088482
hell: 0.11575448122749286
serving: 0.10519131145407251
fast: 0.10306897912311283
food: 0.10257266777068216
dash: 0.1006761861203253
kitchenette: 0.09913501368822292
tasty: 0.09839976762429138
ingredients: 0.09572970581107275
prepared: 0.09283967327745231
pancakes: 0.0902039422284149
game: 0.09010845795632269
frenzy: 0.087328779408056
burgers: 0.08497960033316625
visitors: 0.08497960033316625
craze: 0.0797552584379176
coffee: 0.0771696541516619
queen: 0.0771696541516619
meals: 0.07501719359260098
making: 0.07160655324486731
stand: 0.06559984508286092
cook: 0.0646791097301573
rush: 0.0646791097301573
hot: 0.061893115518301546
serve: 0.0618
```